```python
from numpy.testing import assert_equal
```

## Problem 1
Write a Python function that will count from 1 to an integer `n` (inclusive). For each count, print the counter value and a message separated by a single space. If the counter is even, the message should be `fizz`, `fuzz` if divisible by 3 and `foo` (only) if divisible by 6.

<!-- grade_id: cell-0e7266cd03645e41, points: 0, grade: false, solution: true -->
```python
def count_and_print(n):
    """Count from 1 to n whilst displaying a message
    
    Count from 1 to an integer n (inclusive). For each count, print the 
    counter value and a message separated by a single space. If the counter is
    even, the message should be `fizz`, `fuzz` if divisible by 3 and `foo` 
    (only) if divisible by 6.
    
    Parameters
    ----------
    n : integer
        Count up to this number (inclusive)
        
    Examples
    --------
    >>> count_and_print(12)
    1
    2 fizz
    3 fuzz
    4 fizz
    5
    6 foo
    7 
    8 fizz
    9 fuzz
    10 fizz
    11
    12 foo

    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
%%capture out1
count_and_print(12)
```

```python
%%capture out2
count_and_print(18)
```

<!-- grade_id: cell-5e68bd9f3673ce01, points: 1, grade: true, solution: false -->
```python
assert_equal(
    out1.stdout,
    '1\n'
    '2 fizz\n'
    '3 fuzz\n'
    '4 fizz\n'
    '5\n'
    '6 foo\n'
    '7\n'
    '8 fizz\n'
    '9 fuzz\n'
    '10 fizz\n'
    '11\n'
    '12 foo\n')
assert_equal(
    out2.stdout,
    '1\n'
    '2 fizz\n'
    '3 fuzz\n'
    '4 fizz\n'
    '5\n'
    '6 foo\n'
    '7\n'
    '8 fizz\n'
    '9 fuzz\n'
    '10 fizz\n'
    '11\n'
    '12 foo\n'
    '13\n'
    '14 fizz\n'
    '15 fuzz\n'
    '16 fizz\n'
    '17\n'
    '18 foo\n'
)
```

## Problem 2
Write a function `overlap_interval` that accepts parameters `start1`, `end1`, `start2` and `end2` corresponding to the endpoints of two numeric intervals. It should print the endpoints of their overlaps or print the message `No overlap!` if they don't overlap.

<!-- grade_id: cell-c3a64bdf8b11fb95, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
%%capture out1
overlap_interval(-5, 10, 5, 12)
```

```python
%%capture out2
overlap_interval(-5, 10, 11, 12)
```

```python
%%capture out3
overlap_interval(-5, 10, -10, 12)
```

```python
%%capture out4
overlap_interval(5, 8, -10, 12)
```

```python
%%capture out5
overlap_interval(-12, 12, 8, 10)
```

<!-- grade_id: cell-7812ed25101d9ddf, points: 1, grade: true, solution: false -->
```python
assert_equal(out1.stdout, '5 10\n')
assert_equal(out2.stdout, 'No overlap!\n')
assert_equal(out3.stdout, '-5 10\n')
```

## Problem 3
Write a function that returns the cosine of an input (in radians) using [Maclaurin's approximation](https://en.wikibooks.org/wiki/Trigonometry/Power_Series_for_Cosine_and_Sine). It should use a `while` loop to compute terms of the summation until the last term is smaller than $10^{-15}$.

<!-- grade_id: cell-54183911ca9ecc32, points: 1, grade: true, solution: true -->
```python
def cosine(theta):
    """Return cosine of theta using Maclaurin's approximation
    
    Parameters
    ----------
    theta : float
        Value to compute cosine of, in radians
        
    Returns
    -------
    cos_theta : float
        Cosine of theta
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

## Problem 4
The greatest common divisor (GCD) of $a$ and $b$ is the largest number that divides both of them with no remainder.

One way to find the GCD of two numbers is based on the observation that if $r$ is the remainder when $a$ is divided by $b$, then $gcd(a, b) = gcd(b, r)$. As a base case, we can use $gcd(a, 0) = a$.

Write a function called `gcd` that takes parameters `a` and `b` and returns their greatest common divisor.

<!-- grade_id: cell-c29ca12ec75789d2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2f070dd683636894, points: 1, grade: true, solution: false -->
```python
assert_equal(gcd(4, 2), 2)
assert_equal(gcd(2, 4), 2)
assert_equal(gcd(13, 5), 1)
assert_equal(gcd(20, 24), 4)
```

## Problem 5
Write a function `biased_sum` that returns the sum of positive integers passed to it but numbers that are multiples of the `base` are doubled before being added. By default, `base` is 2. Return `None` if there are no numbers passed to it.

<!-- grade_id: cell-6295767c33976fea, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-269f051cdd06af87, points: 1, grade: true, solution: false -->
```python
assert_equal(biased_sum(1, 2, 3), 8)
assert_equal(biased_sum(1, 2, 3, 4, 5, 6, 3), 36)
assert_equal(biased_sum(1, 2, 3, 4, 5, 6, base=3), 30)
```

## Problem 6
Write a function `last_in_sequence` that accepts a string of digits, which can have a repeating sequence of 0 to 9 in it but may have intervening digits between them. The function should return the last digit in the sequence found in the string. It should return `None` if there is no such sequence in the string.

Examples:
* **0**9**12**251**3**2**4**21 $\rightarrow$ 4
* 1**0**9**12**251**3**2**4**21 $\rightarrow$ 4
* **012345678901234** $\rightarrow$ 4
* 1912251324 $\rightarrow$ `None`

<!-- grade_id: cell-f26b1b122b693409, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-19149d85bde88dea, points: 1, grade: true, solution: false -->
```python
assert_equal(last_in_sequence('0912251324'), 4)
assert_equal(last_in_sequence('10912251324'), 4)
assert_equal(last_in_sequence('012345678901234'), 4)
assert_equal(last_in_sequence('1912251324'), None)
```

## Problem 7
Write a function `check_password` that returns `True` if the input password satisfies the following criteria:

* at least 8 characters in length
* contains at least one lowercase letter
* contains at least one uppercase letter
* contains at least one number

It should return `False` if at least one of the criteria is not satisfied.

```python
def check_password(password):
    if (any(x.isupper() for x in password) and 
        any(x.islower() for x in password) and 
        any(x.isdigit() for x in password) and 
        len(password)>=8):
        return True
    else:
        return False
```

```python
check_password('abC12345')
```

Output:
```
True
```

```python
'password'.isupper()
```

Output:
```
False
```

<!-- grade_id: cell-a72d390a2aa6048b, points: 0, grade: false, solution: true -->
```python
def check_password(password):
    # YOUR CODE HERE
    import string

    numlist = [str(i) for i in range(0,10)]
    lower = string.ascii_lowercase
    upper = string.ascii_uppercase
    alpha_lower = list(lower)
    alpha_upper = list(upper)
    sort_pass = []

    if len(password) < 8:
        return False 
    else:
        pass

    for i in password:       
        sort_pass.append(i)
        sort_pass.sort()

    num = []
    upper2 = []
    lower2 = []
    for i in sort_pass:
   
        if i in numlist:
            num.append(i)
        elif i in alpha_upper:
            upper2.append(i)
        elif i in alpha_lower:
            lower2.append(i)
  
    if len(num) == 0 or len(upper2) == 0 or len(lower2) == 0:
        return False
    else:
        return True

```

```python

```

<!-- grade_id: cell-b8041b3660eec0bd, points: 1, grade: true, solution: false -->
```python
assert_equal(check_password('abcABC1'), False)
assert_equal(check_password('abcABC12'), True)
assert_equal(check_password('abcABCDE'), False)
assert_equal(check_password('abc12345'), False)
assert_equal(check_password('12345ABC'), False)
```

## Problem 8
Write a function `is_palindrome` that returns `True` if the input string is a palindrome and `False` otherwise.

<!-- grade_id: cell-c7e575e3d909fd56, points: 0, grade: false, solution: true -->
```python
def is_palindrome(text):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-71c90ba9148f6663, points: 1, grade: true, solution: false -->
```python
assert_equal(is_palindrome('abcd'), False)
assert_equal(is_palindrome('abba'), True)
assert_equal(is_palindrome('abeba'), True)
assert_equal(is_palindrome('abcda'), False)
assert_equal(is_palindrome('ab cd'), False)
assert_equal(is_palindrome('ab ba'), True)
assert_equal(is_palindrome(' ab ba'), True)
assert_equal(is_palindrome('ab  ba'), True)
```

## Problem 9
Write a function `create_squares` that takes in the number of stars on a side of a square then returns a 2$\times$2 grid of squares as a string. Separate each line by the newline (`\n`) character.

<!-- grade_id: cell-9d0bd755f6457ac0, points: 0, grade: false, solution: true -->
```python
def create_squares(num_stars):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-1f0947cbeb45f780, points: 1, grade: true, solution: false -->
```python
assert_equal(
    create_squares(1),
"""+ - + - +
| * |   |
+ - + - +
|   | * |
+ - + - +
"""
)
assert_equal(
    create_squares(3),
"""+ - - - + - - - +
| * * * |       |
| * * * |       |
| * * * |       |
+ - - - + - - - +
|       | * * * |
|       | * * * |
|       | * * * |
+ - - - + - - - +
"""
)
```

## Problem 10
Write a function `create_grid` that takes in the number squares on a side and the number of stars on a side of a square then returns the grid of squares as a string. Separate each line by the newline (`\n`) character.


<!-- grade_id: cell-d554859dd48da4e8, points: 0, grade: false, solution: true -->
```python
def create_grid(num_squares, num_stars):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-804b324f525cec69, points: 1, grade: true, solution: false -->
```python
assert_equal(
    create_grid(1, 1),
"""+ - +
| * |
+ - +
""")
assert_equal(
    create_grid(2, 4),
"""+ - - - - + - - - - +
| * * * * |         |
| * * * * |         |
| * * * * |         |
| * * * * |         |
+ - - - - + - - - - +
|         | * * * * |
|         | * * * * |
|         | * * * * |
|         | * * * * |
+ - - - - + - - - - +
""")
assert_equal(
    create_grid(4, 2),
"""+ - - + - - + - - + - - +
| * * |     | * * |     |
| * * |     | * * |     |
+ - - + - - + - - + - - +
|     | * * |     | * * |
|     | * * |     | * * |
+ - - + - - + - - + - - +
| * * |     | * * |     |
| * * |     | * * |     |
+ - - + - - + - - + - - +
|     | * * |     | * * |
|     | * * |     | * * |
+ - - + - - + - - + - - +
""")
```